# Exercise: Image Classification

In this exercise, we want to look at image classification using support vector machines. We will first use a simple data set from the hand segmentation to see how the SVM works, and then train a classifier on a larger data set of images.


We will use the SVM implementation from sklearn, which in turn uses the svm library.


```python
import cv2
import os
import sys
import numpy as np
import sklearn

# for histogram of oriented gradients (hog)
from skimage.feature import hog

# we will use the svm from sklearn
from sklearn import svm

sys.path.append('.')
from cifar import load_cifar

# for displaying images in jupyter
import matplotlib as mpl
from matplotlib import pyplot as plt

from pathlib import Path
data_dir = Path('/exchange/cvai/images')
```

## Exercise 1: Hand / Background Segmentation

In the exercise in week 3, we have looked at segmenting the foreground (the hands) from a green background. We will now try to use machine learning for the same task. This is actual a task of semantic segmentation, as the classification is on pixel level which we will look at in more detail later in the lecture.

This is the image that we want to use for training. There is a similar image that we can use for testing.

```python
image  = cv2.imread(str(data_dir / 'hand_small.png'))
image = cv2.cvtColor(image, cv2.COLOR_BGR2RGB)
label = cv2.imread(str(data_dir / 'label_small.png'))
plt.subplot(1, 2, 1)
plt.imshow(image)
plt.subplot(1, 2, 2)
plt.imshow(label)
print(image.shape)
print(label.shape)
```

We would like to use 2D color features which should not contain the brightness. So we will convert the images to HSV and then use only hue and saturation channels. The SVM only needs an array or features, so we have to reshape it.

```python
hsv = cv2.cvtColor(image, cv2.COLOR_RGB2HSV)
hs = hsv[:,:,0:2]
data = hs.reshape((-1,2))
label_grey = label[:,:,0]
classes = label_grey.reshape(-1)
```

As now each pixel is a training example, the dataset is fairly large. We will reduce it by taking only every n-th pixel. You can try to lower n later, but training will take some time.

```python
n = 50
classes_reduced = classes[1::n]
data_reduced = data[1::n,:]
```

We can plot the features (as they are 2D) to see how difficult a classification is.

```python
positives = data_reduced[classes_reduced==0]
negatives = data_reduced[classes_reduced==255]

fig, ax = plt.subplots(figsize=(4, 4))
ax.scatter(positives[:,0], positives[:,1], marker='x', label='Hand (foreground, positive)')
ax.scatter(negatives[:,0], negatives[:,1], marker='o', label='Background (negative)')
ax.set_xlabel('Hue'); ax.set_xlim([0, 255])
ax.set_ylabel('Saturation'); ax.set_ylim([0, 255])
ax.legend()
```

As you can see, it should be quite easy to train a linear classifier on this data set. We will now use the SM

```python
classifier = svm.SVC(C=1.0, kernel='linear', gamma='auto')
```

```python
classifier.fit(data_reduced, classes_reduced)
```

We can apply the classifier to check the result. The score is the mean accuracy.

```python
classifier.score(data_reduced,classes_reduced)
```

Now load the second image and apply the same data reduction and calculate the score in a variable `score`.

<!-- grade_id: cell-ae1f82050010d1d2, points: 0, grade: false, solution: true -->
```python
image2  = cv2.imread(str(data_dir /'hand_small_2.png'))
image2 = cv2.cvtColor(image2, cv2.COLOR_BGR2RGB)
label2 = cv2.imread(str(data_dir / 'label_small_2.png'))
score = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-415c806a03b1ab32, points: 1, grade: true, solution: false -->
```python
np.testing.assert_allclose(score, 0.9, atol=0.1)
```

We have seen, that the score is not quite as good on the test image than on the training image.

If you like, you can also apply the classification to the whole image and then plot the result. However that will take a while to calculate.

## Exercise 2: Image Classification on CIFAR-10

We now want to classify images.

The data set that we use is the CIFAR data set. It contains 60000 small (32x32) images of 10 different classes of which 50000 are in the training set and 10000 in the test set. There is a version of the data set installed in the /home/data folder. We will use that, the utilities for loading the data are in the cifar.py file in the same directory.

```python
cifar_path = data_dir / 'cifar-10-batches-py'
(x_image_train, y_train), (x_image_test, y_test) = load_cifar(cifar_path)
print(x_image_train.shape)
print(x_image_test.shape)
```

The following method will display some of the images and labels.

```python
def show_images(images, labels):
    label_names = ['airplane', 'automobile', 'bird', 'cat', 'deer', 'dog', 'frog', 'horse', 'ship', 'truck']
    fig, axes = plt.subplots(3, 5)
    fig.subplots_adjust(hspace=0.6, wspace=0.3)
    
    for i, ax in enumerate(axes.flat):
        # Plot image.
        example = np.random.randint(len(images))
        ax.imshow(images[example])
        
        # Name of the true class.
        cls_true_name = label_names[labels[example, 0]]
        xlabel = "{0}".format(cls_true_name)
        
        # Show the classes as the label on the x-axis.
        ax.set_xlabel(xlabel)
        
        # Remove ticks from the plot.
        ax.set_xticks([])
        ax.set_yticks([])
    return
```

```python
mpl.rcParams['figure.dpi']= 100
show_images(x_image_train, y_train)
```

We need image features as input. There are several possibilities, for example we could just use a color distribution. One of the best features are HOGs, histograms of oriented gradients another would be SIFT. We will use the HOG features in this example. Let us calculate the feature for one image and display it.

```python
example_image = x_image_train[1]
example_hog, example_hog_image = hog(example_image, channel_axis=2, visualize=True)

_, axes = plt.subplots(1, 2, figsize=(6, 6), sharex=True, sharey=True)
axes[0].imshow(example_image)
axes[1].imshow(example_hog_image)
plt.show()
```

```python
print(example_hog.shape)
print(example_hog)
```

50000 images is a bit large, so we will start with a smaller data set of 10000 images randomly sampled from the full data set.

```python
indices = np.random.choice(x_image_train.shape[0], 10000, replace=False)
x_image_train_small = x_image_train[indices, :, :, :]
y_train_small = y_train[indices, :]
print(x_image_train_small.shape)
print(y_train_small.shape)
```

We will now need the feature vectors of all images, define a function that calculate the hog features on an array of images and returns an array of feature vectors.

<!-- grade_id: cell-3b68a73cd7ae6bd0, points: 0, grade: false, solution: true -->
```python
def calculate_features_hog(images):
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-adc69758f22047bb, points: 0, grade: true, solution: false -->
```python
features_train_small = calculate_features_hog(x_image_train_small)
assert features_train_small.shape[0] == 10000
assert features_train_small.shape[1] == 324
```

Now train a SVM classifier on that data set as above, again use a linear kernel.

<!-- grade_id: cell-1fe082c65f599a96, points: 0, grade: false, solution: true -->
```python
classifier = None

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4adb173546b145c6, points: 1, grade: true, solution: false -->
```python
# we will test the result by applying it to the first image, however 
# we can not really check the result, as it depends on the selection of data
result = classifier.predict([features_train_small[0]])
print(result)
```

### Confusion matrix

In order to evaluate the result, we want to have a look at the confusion matrix which can be calculated using
 `sklearn.metrics.confusion_matrix`. Calculate the prediction and then the confusion matrix using the labels.

<!-- grade_id: cell-9d0795c070990628, points: 0, grade: false, solution: true -->
```python
cm = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b1373c42ee2d934f, points: 1, grade: true, solution: false -->
```python
assert cm.shape[0] == 10
assert cm.shape[1] == 10
```

```python
# there is also a display function for the matrix
disp = sklearn.metrics.ConfusionMatrixDisplay(cm)
disp.plot()
```

For an overall assessment of the classifier, there is the function `sklearn.metrics.classification_report`.

```python
label_names = ['airplane', 'automobile', 'bird', 'cat', 'deer', 'dog', 'frog', 'horse', 'ship', 'truck']
print(sklearn.metrics.classification_report(y_train_small, y_train_small_predict, target_names=label_names))
```

### Optional further exercises
Of course we should evaluate the results on the test set instead, how does it look there? The results are not that good, can you change any parameters of the SVM to get better results?

```python

```