# Introduction to Segmentation: Clustering algorithms


One of the basic unsupervised learning approaches is clustering, where the algorithm has to find out which part of the data belongs together in one cluster. 

For segmentation, this data is taken from the pixel values, possible from a neighborhood of the pixels and often by calculating specific features in this region. However, for the moment we will use the raw (color) data of the pixels and possibly also the location and look at features later.

```python
import numpy as np
import skimage
import skimage.data
import skimage.io
import skimage.future
import skimage.segmentation
import sklearn.cluster
import cv2

# for displaying images in jupyter
from matplotlib import pyplot as plt

%matplotlib inline
import matplotlib.pyplot as plt
plt.rcParams['figure.figsize'] = [10, 10]
#mpl.rcParams['figure.dpi']= 400
```

## Exercise 1: Using KMeans for background subtraction on a medical image

In the first exercise, we want to extract the background from an image showning the hands on green background.

<img src=hand_small.png width=400px>

This could also be achieved by defining a suitable range of color for either the background or the foreground and use thresholding. We want to find out, if it is also possible using k-Means.

Use kmeans for clustering all the pixels in the image into 2 clusters using only the RGB values.

If we use the raw pixels, the position of the pixels do not matter. You will have to reshape the image into a 1D Array of RGB values and convert them to float to be able to use the library function.

There are different implementations of kmeans available, most noteworthy an opencv and a sklearn implementation. We will use the sklearn implementation:

https://scikit-learn.org/stable/modules/generated/sklearn.cluster.KMeans.html


```python
image  = skimage.io.imread('/exchange/cvai/images/hand_small.png')
plt.imshow(image)
```

### k-Means
Use k-means to find 2 clusters and display them.

For displaying you can use either the found cluster centers as colors or overlay the labels, see https://scikit-image.org/docs/dev/api/skimage.color.html#skimage.color.label2rgb

<!-- grade_id: cell-a3c4b1b388fcf880, points: 0, grade: false, solution: true -->
```python
# The first step should be to generate a 1D array of type float32 of the data
data = None
# YOUR CODE HERE
raise NotImplementedError()

# then create a kmeans object and fit the data
kmeans = None
# YOUR CODE HERE
raise NotImplementedError()

# you can check the cluster centers from the kmeans object, this should be close to the color 
# of the hand and of the background.
center_colors = np.uint8(kmeans.cluster_centers_)

# try to plot the result s an image again, there are different ways to do this...

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2d4d04aeed13a0f7, points: 2, grade: true, solution: false -->
```python
np.testing.assert_array_equal(data.shape, (497664, 3))
np.testing.assert_array_equal(kmeans.cluster_centers_.shape, (2,3))
# reshape the class labels
label_image_test = kmeans.labels_.reshape((image.shape[0], image.shape[1]))
# these two points are clearly on the background and should have the same label
assert label_image_test[100,100] == label_image_test[200,400]
# this point is on a hand and should have a different label
assert label_image_test[100,100] != label_image_test[400, 200]
```

How can you find out, if the segmentation was good or not? I.e. if you run it twice, which result should you keep?


<!-- grade_id: cell-28c978f118851ffe, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Exercise 2: K-Means with multiple labels.

Now try the K-Means Algorithmus on the following image to separate the different colored figures and the gray and black background. Ideally we want to have a cluster for every colored figure, one for the table and one for the background, that would be 6 classes.

```python
image = skimage.io.imread('/exchange/cvai/images/carcassonne_figures.jpg')

plt.rcParams['figure.figsize'] = [10, 10]
plt.imshow(image)
```

<!-- grade_id: cell-3aadfb969ed739e7, points: 0, grade: false, solution: true -->
```python
kmeans = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-5b5453ce65e0ccd5, points: 2, grade: true, solution: false -->
```python
# results depend a bit on the initial conditions of the center, so we dont do any real checks here
assert kmeans is not None
assert kmeans.cluster_centers_.shape[0] == 6
```

Using 6 classes probably did not generate a very good result. Try different values for the number of clusters. What results do you get? What is the best parameter setting.

<!-- grade_id: cell-ba80452560ef68b4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

## Exercise 3: Using different color spaces

One of the problems is that similar colors with different intensity will end up in different clusters. How could you solve that issue?

One idea would be to use a color space that has the luminance (or value) as a seperate channel, and then only use the other 2 channels. A possible color space is Lab, but you could also try different ones. You will find color space transfer routines in skimage.color.

Convert the image into another color space and display the 3 channels as gray value image.

<!-- grade_id: cell-f1e077038c955c50, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now we want to drop one of the color spaces and reshape the data so that we have an array of 2 channels instead of 3. Save this data in the variable `data`

<!-- grade_id: cell-b7de6be628ac9f21, points: 0, grade: false, solution: true -->
```python
data = None
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-15af7c1082b5236e, points: 2, grade: true, solution: false -->
```python
np.testing.assert_array_equal(data.shape, (623808, 2))
```

One of the advantages of 2D data is that we can plot it direclty:

```python
plt.rcParams['figure.figsize'] = [12, 12]
plt.plot(data[:,0], data[:,1], '.')
np.max(data[:,0])
np.max(data[:,1])
```

Now try to cluster this data into 5 clusters (as background and table are both grey, respectively black). Does it work better?


<!-- grade_id: cell-b85e0241c35768a3, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

We can also display the clustering the feature space.

```python
plt.scatter(x=data[:,0], y=data[:,1], c=kmeans.labels_)
```

## Exercise 4 (optional): k-Means with position and color

Implement a k-Means clustering that uses the position of a pixel together with the color for clustering, does this work better?


<!-- grade_id: cell-cc98c46f4b5b9cd6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

```

## Exercise 5: Mean Shift

As the last exercise, we would like to solve the same clustering using mean shift. There is an implementation in sklearn:

https://scikit-learn.org/stable/modules/generated/sklearn.cluster.MeanShift.html

Mean shift is quite a slow algorithm, so we will try with a reduced size image first.


```python
image = skimage.io.imread('/exchange/cvai/images/carcassonne_figures.jpg')
# this reduces the image size by an integer factor using the mean of the neighborhood
image_small = skimage.transform.downscale_local_mean(image, factors=(2,2,1)).astype(int)
data = image_small.reshape((-1,3)).astype(np.float32)
plt.imshow(image_small)
```

The mean shift algorithms will first evaluate the kernel to be used. This can be done seperately to have more control over values.

```python
bandwidth = sklearn.cluster.estimate_bandwidth(data, quantile=0.2, n_samples=5000)
print(bandwidth)
```

Now create a meanshift object, fit the data and display the result again.

<!-- grade_id: cell-57995b92b3e382a2, points: 0, grade: false, solution: true -->
```python
ms = None
# create sklearn.cluster.MeanShift object, fit the data and display the result
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-10e696ff6f630638, points: 2, grade: true, solution: false -->
```python
assert ms.cluster_centers_.shape[0] > 10
```

This is all for the exercise.

If you feel like implementing more, how about programming a region growing algorithm :-).